For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

# On se remet à Python

## Des exercices de bases


### Des calculs simples

<!-- grade_id: cell-8624d910f5d99547, points: 0, grade: false, solution: true -->
```python
def moyenne_3_nb(n1,n2,n3):
    """Compute the average of 3 numbers
    """
    ### BEGIN SOLUTION
    return (n1+n2+n3)/3
    ### END SOLUTION

```

<!-- grade_id: cell-2e156d82f5ec7bc1, points: 2, grade: true, solution: false -->
```python
""" Check your function """ 
assert moyenne_3_nb(10,10,10) == 10
assert moyenne_3_nb(10,15,20) == 15
assert moyenne_3_nb(-5,10,16) == 7
```

<!-- grade_id: cell-58a579236231e23e, points: 0, grade: false, solution: true -->
```python
def moyenne_3_nb_pondere(n1,n2,n3,p1,p2,p3):
    """Compute the average of 3 numbers with coefficients 
        raise an error if the sum of the coefficient is equal to 0
    """
    ### BEGIN SOLUTION
    if(p1+p2+p3 == 0):
        raise ValueError("division by 0")
    return ((n1*p1)+(n2*p2)+(n3*p3))/(p1+p2+p3)
    ### END SOLUTION
```

<!-- grade_id: cell-e12e4f9f4ca6bb9f, points: 2, grade: true, solution: false -->
```python
try :
    moyenne_3_nb_pondere(1,2,4,2,2,-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

assert moyenne_3_nb_pondere(1,2,4,2,2,1) == 2
assert moyenne_3_nb_pondere(0,0,0,2,2,3) == 0
```

### Des conditions et des boucles

<!-- grade_id: cell-1d2dbf883de032be, points: 0, grade: false, solution: true -->
```python
def sum_odd(n):
    """Sum of odd numbers less or equal to n"""
    ### BEGIN SOLUTION
    sum = 0
    for i in range(1,n+1,2):
        sum = sum + i
    return sum
    ### END SOLUTION
```

<!-- grade_id: cell-7da3d03455537762, points: 3, grade: true, solution: false -->
```python
assert sum_odd(0) == 0
assert sum_odd(1) == 1
assert sum_odd(2) == 1
assert sum_odd(3) == 4
assert sum_odd(20) == 100
```

<!-- grade_id: cell-1d2dbf883de032bf, points: 0, grade: false, solution: true -->
```python
def sum_first_odd(n):
    """Sum of n first odd numbers """
    ### BEGIN SOLUTION
    sum = 0
    nb = 0
    i = 1
    while nb < n : 
        sum = sum + i
        nb = nb + 1
        i = i + 2
    return sum
    ### END SOLUTION
```

<!-- grade_id: cell-7da3d0345553776a, points: 4, grade: true, solution: false -->
```python
assert sum_first_odd(0) == 0
assert sum_first_odd(1) == 1
assert sum_first_odd(2) == 4
assert sum_first_odd(3) == 9
assert sum_first_odd(20) == 400
```

## Des chaînes de caractères

### Supression

Ecrire la fonction `suppression(c, s)` qui supprime toutes les occurrences du caractère `c` dans la chaîne `s`.

<!-- grade_id: cell-5388e7c7c7ee2ea2, points: 0, grade: false, solution: true -->
```python
def suppression(c,s):
    ### BEGIN SOLUTION
    new_s = ""
    for a in s:
        if a != c:
            new_s = new_s + a
    return new_s            
    ### END SOLUTION
```

Output:
```
'ppl'
```

<!-- grade_id: cell-021289affc28c89d, points: 2, grade: true, solution: false -->
```python
assert suppression('a',"youplapla") == "youplpl"
assert suppression("z","polytech")== "polytech"
assert suppression('o',"oooooo") == ""
```

### Supression 2

Supprimer la première occurence d'une lettre dans un mot

<!-- grade_id: cell-5388e7c7c7ee2ef, points: 0, grade: false, solution: true -->
```python
def suppression_premiere(c,s):
    ### BEGIN SOLUTION
    new_s = ""
    vu = False
    for a in s:
        if a != c or vu:
            new_s = new_s + a
        else : 
            vu = True
    return new_s            
    ### END SOLUTION
```

Output:
```
'youplplapla'
```

<!-- grade_id: cell-021289affc28c89e, points: 2, grade: true, solution: false -->
```python
assert suppression_premiere('a',"youplaplapla") == "youplplapla"
assert suppression_premiere("z","polytech")== "polytech"
assert suppression_premiere('o',"oooooo") == "ooooo"
```

### Palindromes

Ecrire la fonction `est_palindrome(s)` qui retourne vraie si la chaîne de caractère `s` est un palindrome

<!-- grade_id: cell-372eac3e065d97df, points: 0, grade: false, solution: true -->
```python
def est_palindrome(s):
    ### BEGIN SOLUTION
    return s == s[::-1] 
    ### END SOLUTION
```

Output:
```
False
```

<!-- grade_id: cell-d3913fe6af2fb45e, points: 3, grade: true, solution: false -->
```python
assert est_palindrome("gag")
assert est_palindrome("laval")
assert est_palindrome("")
assert est_palindrome("palindrome") == False
```

### Miroir

Ecrire la fonction miroir qui crée un palindrome. Elle prend un chaine en paramètre et retourne le palindrome correspondant au miroir de la chaine

<!-- grade_id: cell-f23ec13cae6ca407, points: 0, grade: false, solution: true -->
```python
def miroir(s):
    ### BEGIN SOLUTION
    return s + s[::-1]
    ### END SOLUTION

```

<!-- grade_id: cell-c639250c614271e5, points: 3, grade: true, solution: false -->
```python
assert est_palindrome(miroir("palindrome"))
assert est_palindrome(miroir(""))
assert est_palindrome(miroir("porte"))
```

## Listes

### Max d'une liste

Ecrire une fonction qui retourne le plus grand nombre d'une liste

<!-- grade_id: cell-7d981642e463b138, points: 0, grade: false, solution: true -->
```python
def max_liste(l):
    max = l[0]
    for i in l[1:]:
        if i > max:
            max = i
    return max
```

<!-- grade_id: cell-5c7267b8c5ff8910, points: 2, grade: true, solution: false -->
```python
assert max_liste([3, 7, 9, 5.4, 8.9, 9, 8.999, 5]) == 9
assert max_liste([-3,5,-22,4,12,2]) == 12
```

### Nombre d'occurence 

Ecrire la fonction qui compte le nombre d'occurence d'un élément dans une liste

<!-- grade_id: cell-3d59f20935e1a8f4, points: 0, grade: false, solution: true -->
```python
def nb_occurences(l,x):
    cpt = 0
    for i in l :
        if i == x:
            cpt = cpt+1
    return cpt
```

<!-- grade_id: cell-99f002006294a3a6, points: 2, grade: true, solution: false -->
```python
assert nb_occurences([3, 7, 9, 5.4, 8.9, 9, 8.999, 5], 9) == 2
assert nb_occurences(["chat", "ours", "chat", "chat", "loup"], "chat") == 3
assert nb_occurences(["chat", "ours", "chat", "chat", "loup"], "chien") == 0
```